# Aula 11 - Análise de Séries Temporais: Prática Independente

**Programação para Ciência de Dados**

**Data:** 11 de Novembro de 2025

---

## Instruções

Este notebook contém exercícios práticos sobre análise de séries temporais.

**Importante:**
- Execute cada célula em ordem
- Não modifique as células de teste
- Suas respostas devem estar nas células indicadas com `# SEU CÓDIGO AQUI`
- Execute os testes para verificar suas respostas

---

## Dataset

Utilizaremos o mesmo dataset Air Quality da aula:
- Dados de qualidade do ar (cidade italiana)
- Medições horárias de março 2004 a fevereiro 2005
- Variáveis: CO, NO2, temperatura, umidade

---

## Setup

```python
# === CONFIGURAÇÃO INICIAL ===

!pip install --upgrade pip --quiet
!pip cache purge
!pip install otter-grader --no-cache-dir -q
!mkdir -p tests

print("Ambiente configurado!")
```

Output:
```
[2K   [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m1.8/1.8 MB[0m [31m14.5 MB/s[0m eta [36m0:00:00[0m
[?25hFiles removed: 6 (1.8 MB)
Ambiente configurado!
```

```python
%%writefile tests/q1.py
OK_FORMAT = True

test = {
    "name": "q1",
    "points": 8,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> import pandas as pd
                    >>> # Verificar se df existe e tem DatetimeIndex
                    >>> assert 'df' in globals(), "DataFrame 'df' não foi criado"
                    >>> assert isinstance(df.index, pd.DatetimeIndex), "O índice deve ser DatetimeIndex"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar se está ordenado
                    >>> assert df.index.is_monotonic_increasing, "O índice deve estar ordenado cronologicamente"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar se não tem duplicatas
                    >>> assert not df.index.has_duplicates, "O índice não deve ter timestamps duplicados"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tamanho do DataFrame
                    >>> assert len(df) > 9000, "DataFrame deve ter mais de 9000 registros"
                    >>> assert len(df) < 10000, "DataFrame deve ter menos de 10000 registros"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q1.py
```

```python
%%writefile tests/q2.py
OK_FORMAT = True

test = {
    "name": "q2",
    "points": 6,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> # Verificar se colunas existem
                    >>> assert 'hour' in df.columns, "Coluna 'hour' não foi criada"
                    >>> assert 'dayofweek' in df.columns, "Coluna 'dayofweek' não foi criada"
                    >>> assert 'month' in df.columns, "Coluna 'month' não foi criada"
                    >>> assert 'is_weekend' in df.columns, "Coluna 'is_weekend' não foi criada"
                    >>> assert 'is_rush_hour' in df.columns, "Coluna 'is_rush_hour' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar valores de hour
                    >>> assert df['hour'].min() >= 0, "hour mínimo deve ser >= 0"
                    >>> assert df['hour'].max() <= 23, "hour máximo deve ser <= 23"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar valores de dayofweek
                    >>> assert df['dayofweek'].min() >= 0, "dayofweek mínimo deve ser >= 0"
                    >>> assert df['dayofweek'].max() <= 6, "dayofweek máximo deve ser <= 6"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar valores binários
                    >>> assert set(df['is_weekend'].unique()).issubset({0, 1}), "is_weekend deve conter apenas 0 ou 1"
                    >>> assert set(df['is_rush_hour'].unique()).issubset({0, 1}), "is_rush_hour deve conter apenas 0 ou 1"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q2.py
```

```python
%%writefile tests/q3.py
OK_FORMAT = True

test = {
    "name": "q3",
    "points": 8,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> import numpy as np
                    >>> # Verificar se variáveis existem
                    >>> assert 'co_mean_weekday' in globals(), "Variável 'co_mean_weekday' não foi criada"
                    >>> assert 'co_mean_weekend' in globals(), "Variável 'co_mean_weekend' não foi criada"
                    >>> assert 'co_mean_rushhour' in globals(), "Variável 'co_mean_rushhour' não foi criada"
                    >>> assert 'co_mean_by_hour' in globals(), "Variável 'co_mean_by_hour' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tipos
                    >>> assert isinstance(co_mean_weekday, (int, float, np.number)), "co_mean_weekday deve ser um número"
                    >>> assert isinstance(co_mean_weekend, (int, float, np.number)), "co_mean_weekend deve ser um número"
                    >>> assert isinstance(co_mean_rushhour, (int, float, np.number)), "co_mean_rushhour deve ser um número"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> import pandas as pd
                    >>> # Verificar co_mean_by_hour
                    >>> assert isinstance(co_mean_by_hour, pd.Series), "co_mean_by_hour deve ser uma Series"
                    >>> assert len(co_mean_by_hour) == 24, "co_mean_by_hour deve ter 24 valores (uma para cada hora)"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar valores positivos
                    >>> assert co_mean_weekday > 0, "co_mean_weekday deve ser positivo"
                    >>> assert co_mean_weekend > 0, "co_mean_weekend deve ser positivo"
                    >>> assert co_mean_rushhour > 0, "co_mean_rushhour deve ser positivo"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q3.py
```

```python
%%writefile tests/q4.py
OK_FORMAT = True

test = {
    "name": "q4",
    "points": 8,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> import pandas as pd
                    >>> # Verificar se variáveis existem
                    >>> assert 'co_daily' in globals(), "Variável 'co_daily' não foi criada"
                    >>> assert 'co_weekly' in globals(), "Variável 'co_weekly' não foi criada"
                    >>> assert 'co_daily_max' in globals(), "Variável 'co_daily_max' não foi criada"
                    >>> assert 'co_6h' in globals(), "Variável 'co_6h' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tipos
                    >>> assert isinstance(co_daily, pd.Series), "co_daily deve ser uma Series"
                    >>> assert isinstance(co_weekly, pd.Series), "co_weekly deve ser uma Series"
                    >>> assert isinstance(co_daily_max, pd.Series), "co_daily_max deve ser uma Series"
                    >>> assert isinstance(co_6h, pd.Series), "co_6h deve ser uma Series"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tamanhos
                    >>> assert len(co_daily) > 300, "co_daily deve ter mais de 300 registros"
                    >>> assert len(co_daily) < 400, "co_daily deve ter menos de 400 registros"
                    >>> assert len(co_weekly) < len(co_daily), "co_weekly deve ter menos registros que co_daily"
                    >>> assert len(co_6h) > len(co_daily), "co_6h deve ter mais registros que co_daily"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar se índices são DatetimeIndex
                    >>> assert isinstance(co_daily.index, pd.DatetimeIndex), "Índice de co_daily deve ser DatetimeIndex"
                    >>> assert isinstance(co_weekly.index, pd.DatetimeIndex), "Índice de co_weekly deve ser DatetimeIndex"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q4.py
```

```python
%%writefile tests/q5.py
OK_FORMAT = True

test = {
    "name": "q5",
    "points": 8,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> # Verificar se colunas existem
                    >>> assert 'CO_MA24' in df.columns, "Coluna 'CO_MA24' não foi criada"
                    >>> assert 'CO_MA168' in df.columns, "Coluna 'CO_MA168' não foi criada"
                    >>> assert 'CO_std24' in df.columns, "Coluna 'CO_std24' não foi criada"
                    >>> assert 'T_MA24' in df.columns, "Coluna 'T_MA24' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar que não estão todas NaN
                    >>> assert df['CO_MA24'].notna().sum() > 0, "CO_MA24 não deve estar completamente NaN"
                    >>> assert df['CO_MA168'].notna().sum() > 0, "CO_MA168 não deve estar completamente NaN"
                    >>> assert df['CO_std24'].notna().sum() > 0, "CO_std24 não deve estar completamente NaN"
                    >>> assert df['T_MA24'].notna().sum() > 0, "T_MA24 não deve estar completamente NaN"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar que valores fazem sentido
                    >>> assert df['CO_MA24'].max() > 0, "CO_MA24 deve ter valores positivos"
                    >>> assert df['CO_std24'].max() > 0, "CO_std24 deve ter valores positivos"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q5.py
```

```python
%%writefile tests/q6.py
OK_FORMAT = True

test = {
    "name": "q6",
    "points": 8,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> # Verificar se colunas existem
                    >>> assert 'CO_lag1' in df.columns, "Coluna 'CO_lag1' não foi criada"
                    >>> assert 'CO_lag24' in df.columns, "Coluna 'CO_lag24' não foi criada"
                    >>> assert 'CO_diff1' in df.columns, "Coluna 'CO_diff1' não foi criada"
                    >>> assert 'CO_diff24' in df.columns, "Coluna 'CO_diff24' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> import pandas as pd
                    >>> # Verificar que primeiro valor de lag1 é NaN
                    >>> assert pd.isna(df['CO_lag1'].iloc[0]), "Primeiro valor de CO_lag1 deve ser NaN"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar que lags têm valores não-NaN
                    >>> assert df['CO_lag1'].notna().sum() > 0, "CO_lag1 deve ter valores não-NaN"
                    >>> assert df['CO_lag24'].notna().sum() > 0, "CO_lag24 deve ter valores não-NaN"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar que diferenças têm valores não-NaN
                    >>> assert df['CO_diff1'].notna().sum() > 0, "CO_diff1 deve ter valores não-NaN"
                    >>> assert df['CO_diff24'].notna().sum() > 0, "CO_diff24 deve ter valores não-NaN"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q6.py
```

```python
%%writefile tests/q7.py
OK_FORMAT = True

test = {
    "name": "q7",
    "points": 8,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> import pandas as pd
                    >>> # Verificar se variáveis existem
                    >>> assert 'co_daily_clean' in globals(), "Variável 'co_daily_clean' não foi criada"
                    >>> assert 'trend' in globals(), "Variável 'trend' não foi criada"
                    >>> assert 'seasonal' in globals(), "Variável 'seasonal' não foi criada"
                    >>> assert 'residual' in globals(), "Variável 'residual' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tipos
                    >>> assert isinstance(co_daily_clean, pd.Series), "co_daily_clean deve ser uma Series"
                    >>> assert isinstance(trend, pd.Series), "trend deve ser uma Series"
                    >>> assert isinstance(seasonal, pd.Series), "seasonal deve ser uma Series"
                    >>> assert isinstance(residual, pd.Series), "residual deve ser uma Series"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tamanhos
                    >>> assert len(co_daily_clean) > 300, "co_daily_clean deve ter mais de 300 dias"
                    >>> assert len(co_daily_clean) < 400, "co_daily_clean deve ter menos de 400 dias"
                    >>> assert co_daily_clean.isna().sum() == 0, "co_daily_clean não deve ter valores NaN"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar que componentes têm mesmo tamanho
                    >>> assert len(trend) == len(co_daily_clean), "trend deve ter mesmo tamanho que co_daily_clean"
                    >>> assert len(seasonal) == len(co_daily_clean), "seasonal deve ter mesmo tamanho que co_daily_clean"
                    >>> assert len(residual) == len(co_daily_clean), "residual deve ter mesmo tamanho que co_daily_clean"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q7.py
```

```python
%%writefile tests/q8.py
OK_FORMAT = True

test = {
    "name": "q8",
    "points": 6,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> import numpy as np
                    >>> # Verificar se variáveis existem
                    >>> assert 'trend_change_pct' in globals(), "Variável 'trend_change_pct' não foi criada"
                    >>> assert 'seasonal_amplitude' in globals(), "Variável 'seasonal_amplitude' não foi criada"
                    >>> assert 'residual_std' in globals(), "Variável 'residual_std' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tipos
                    >>> assert isinstance(trend_change_pct, (int, float, np.number)), "trend_change_pct deve ser um número"
                    >>> assert isinstance(seasonal_amplitude, (int, float, np.number)), "seasonal_amplitude deve ser um número"
                    >>> assert isinstance(residual_std, (int, float, np.number)), "residual_std deve ser um número"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar valores razoáveis
                    >>> assert -100 < trend_change_pct < 100, "trend_change_pct deve estar entre -100 e 100"
                    >>> assert seasonal_amplitude > 0, "seasonal_amplitude deve ser positivo"
                    >>> assert residual_std > 0, "residual_std deve ser positivo"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q8.py
```

```python
%%writefile tests/q9.py
OK_FORMAT = True

test = {
    "name": "q9",
    "points": 8,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> import pandas as pd
                    >>> import numpy as np
                    >>> # Verificar se variáveis existem
                    >>> assert 'rolling_mean' in globals(), "Variável 'rolling_mean' não foi criada"
                    >>> assert 'rolling_std' in globals(), "Variável 'rolling_std' não foi criada"
                    >>> assert 'zscore' in globals(), "Variável 'zscore' não foi criada"
                    >>> assert 'is_anomaly' in globals(), "Variável 'is_anomaly' não foi criada"
                    >>> assert 'num_anomalies' in globals(), "Variável 'num_anomalies' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tipos
                    >>> assert isinstance(rolling_mean, pd.Series), "rolling_mean deve ser uma Series"
                    >>> assert isinstance(rolling_std, pd.Series), "rolling_std deve ser uma Series"
                    >>> assert isinstance(zscore, pd.Series), "zscore deve ser uma Series"
                    >>> assert isinstance(is_anomaly, pd.Series), "is_anomaly deve ser uma Series"
                    >>> assert isinstance(num_anomalies, (int, np.integer)), "num_anomalies deve ser um inteiro"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar valores de is_anomaly
                    >>> assert set(is_anomaly.dropna().unique()).issubset({0, 1}), "is_anomaly deve conter apenas 0 ou 1"
                    >>> assert num_anomalies >= 0, "num_anomalies deve ser não-negativo"
                    >>> assert num_anomalies == is_anomaly.sum(), "num_anomalies deve ser igual à soma de is_anomaly"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q9.py
```

```python
%%writefile tests/q10.py
OK_FORMAT = True

test = {
    "name": "q10",
    "points": 8,
    "suites": [
        {
            "cases": [
                {
                    "code": r"""
                    >>> import numpy as np
                    >>> # Verificar se variáveis existem
                    >>> assert 'worst_hour' in globals(), "Variável 'worst_hour' não foi criada"
                    >>> assert 'worst_day' in globals(), "Variável 'worst_day' não foi criada"
                    >>> assert 'weekend_vs_weekday_ratio' in globals(), "Variável 'weekend_vs_weekday_ratio' não foi criada"
                    >>> assert 'trend_direction' in globals(), "Variável 'trend_direction' não foi criada"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar tipos (aceita int Python ou numpy)
                    >>> assert isinstance(worst_hour, (int, np.integer)) or isinstance(int(worst_hour), int), "worst_hour deve ser um inteiro"
                    >>> assert isinstance(worst_day, (int, np.integer)) or isinstance(int(worst_day), int), "worst_day deve ser um inteiro"
                    >>> assert isinstance(weekend_vs_weekday_ratio, (int, float, np.number)), "weekend_vs_weekday_ratio deve ser um número"
                    >>> assert isinstance(trend_direction, str), "trend_direction deve ser uma string"
                    """,
                    "hidden": False,
                    "locked": False
                },
                {
                    "code": r"""
                    >>> # Verificar valores válidos
                    >>> assert 0 <= int(worst_hour) <= 23, "worst_hour deve estar entre 0 e 23"
                    >>> assert 0 <= int(worst_day) <= 6, "worst_day deve estar entre 0 e 6"
                    >>> assert float(weekend_vs_weekday_ratio) > 0, "weekend_vs_weekday_ratio deve ser positivo"
                    >>> assert trend_direction in ['up', 'down', 'stable'], "trend_direction deve ser 'up', 'down' ou 'stable'"
                    """,
                    "hidden": False,
                    "locked": False
                }
            ],
            "scored": True,
            "setup": "",
            "teardown": "",
            "type": "doctest"
        }
    ]
}
```

Output:
```
Writing tests/q10.py
```

```python
# Imports
import otter
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
from datetime import datetime, timedelta
import zipfile
import io
import requests
import warnings
warnings.filterwarnings('ignore')

grader = otter.Notebook()

# Configurações
pd.set_option('display.max_columns', None)
plt.style.use('seaborn-v0_8-darkgrid')
plt.rcParams['figure.figsize'] = (12, 5)

print("Setup completo!")
print("✓ Otter Grader carregado")
print("✓ Pandas importado")
print(f"✓ Versão Pandas: {pd.__version__}")
```

Output:
```
Setup completo!
✓ Otter Grader carregado
✓ Pandas importado
✓ Versão Pandas: 2.2.2
```

## Carregar Dados

Execute a célula abaixo para carregar o dataset.

```python
# Carregar dados
zip_url = "https://archive.ics.uci.edu/ml/machine-learning-databases/00360/AirQualityUCI.zip"
r = requests.get(zip_url)
z = zipfile.ZipFile(io.BytesIO(r.content))

with z.open('AirQualityUCI.csv') as f:
    df_raw = pd.read_csv(f, sep=';', decimal=',', na_values=-200)

print(f"Dados carregados: {df_raw.shape}")
df_raw.head()
```

Output:
```
Dados carregados: (9471, 17)
         Date      Time  CO(GT)  PT08.S1(CO)  NMHC(GT)  C6H6(GT)  \
0  10/03/2004  18.00.00     2.6       1360.0     150.0      11.9   
1  10/03/2004  19.00.00     2.0       1292.0     112.0       9.4   
2  10/03/2004  20.00.00     2.2       1402.0      88.0       9.0   
3  10/03/2004  21.00.00     2.2       1376.0      80.0       9.2   
4  10/03/2004  22.00.00     1.6       1272.0      51.0       6.5   

   PT08.S2(NMHC)  NOx(GT)  PT08.S3(NOx)  NO2(GT)  PT08.S4(NO2)  PT08.S5(O3)  \
0         1046.0    166.0        1056.0    113.0        1692.0       1268.0   
1          955.0    103.0        1174.0     92.0        1559.0        972.0   
2          939.0    131.0        1140.0    114.0        1555.0       1074.0   
3          948.0    172.0        1092.0    122.0        1584.0       1203.0   
4          836.0    131.0        1205.0    116.0        1490.0       1110.0   

      T    RH      AH  Unnamed: 15  Unnamed: 16  
0  13.6  48.9  0.7578          NaN   
```

---

## Exercício 1: Configurar DatetimeIndex

Crie uma coluna `timestamp` combinando as colunas `Date` e `Time`, depois defina-a como índice do DataFrame.

**Passos:**
1. Combine `Date` e `Time` em uma nova coluna `timestamp` usando `pd.to_datetime()`
   - Use o formato: `'%d/%m/%Y %H.%M.%S'`
   - Use `dayfirst=True` e `errors='coerce'`
2. Defina `timestamp` como índice usando `.set_index()`
3. Ordene o DataFrame pelo índice usando `.sort_index()`
4. Remova duplicatas do índice usando `~df.index.duplicated(keep='first')`
5. Armazene o resultado em `df`

```python
df_raw['timestamp'] = pd.to_datetime(
    df_raw['Date'] + ' ' + df_raw['Time'],
    format='%d/%m/%Y %H.%M.%S',
    dayfirst=True,
    errors='coerce'
)

df = df_raw.set_index('timestamp').sort_index()
df = df[~df.index.duplicated(keep='first')]

# Remova apenas linhas cujo índice é NaT (timestamp inválido)
df = df[~df.index.isna()]
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

---

## Exercício 2: Extrair Componentes Temporais

Extraia componentes temporais do índice e crie novas colunas.

**Crie as seguintes colunas:**
- `hour`: hora do dia (0-23)
- `dayofweek`: dia da semana (0=Segunda, 6=Domingo)
- `month`: mês (1-12)
- `is_weekend`: 1 se fim de semana (sábado ou domingo), 0 caso contrário
- `is_rush_hour`: 1 se hora de pico (7, 8, 9, 17, 18 ou 19), 0 caso contrário

```python
df['hour'] = df.index.hour
df['dayofweek'] = df.index.dayofweek
df['month'] = df.index.month
df['is_weekend'] = df['dayofweek'].isin([5, 6]).astype(int)
df['is_rush_hour'] = df['hour'].isin([7, 8, 9, 17, 18, 19]).astype(int)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

---

## Exercício 3: Análise por Período

Calcule estatísticas de CO por diferentes períodos.

**Calcule:**
1. `co_mean_weekday`: média de CO em dias úteis (dayofweek < 5)
2. `co_mean_weekend`: média de CO em fins de semana (dayofweek >= 5)
3. `co_mean_rushhour`: média de CO em horários de pico (is_rush_hour == 1)
4. `co_mean_by_hour`: Series com média de CO por hora do dia (use groupby)

**Dica:** Use `.mean()` para calcular médias e ignore valores NaN.

```python
co_mean_weekday = df.loc[df['dayofweek'] < 5, 'CO(GT)'].mean()
co_mean_weekend = df.loc[df['dayofweek'] >= 5, 'CO(GT)'].mean()
co_mean_rushhour = df.loc[df['is_rush_hour'] == 1, 'CO(GT)'].mean()
co_mean_by_hour = df.groupby('hour')['CO(GT)'].mean()
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

---

## Exercício 4: Resample para Diferentes Frequências

Faça resample da coluna `CO(GT)` para diferentes frequências temporais.

**Crie:**
1. `co_daily`: resample diário com média
2. `co_weekly`: resample semanal com média
3. `co_daily_max`: resample diário com máximo
4. `co_6h`: resample a cada 6 horas com média

**Dica:** Use `.resample()` com as frequências: 'D' (dia), 'W' (semana), '6H' (6 horas)

```python
co_daily = df['CO(GT)'].resample('D').mean()
co_weekly = df['CO(GT)'].resample('W').mean()
co_daily_max = df['CO(GT)'].resample('D').max()
co_6h = df['CO(GT)'].resample('6H').mean()
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

---

## Exercício 5: Rolling Windows

Crie médias móveis e outras estatísticas usando rolling windows.

**Crie as seguintes colunas no DataFrame `df`:**
1. `CO_MA24`: média móvel de 24 horas de CO
2. `CO_MA168`: média móvel de 168 horas (7 dias) de CO
3. `CO_std24`: desvio padrão móvel de 24 horas de CO
4. `T_MA24`: média móvel de 24 horas de temperatura (coluna 'T')

**Dica:** Use `.rolling(window, min_periods=1)` seguido de `.mean()` ou `.std()`

```python
df['CO_MA24'] = df['CO(GT)'].rolling(window=24, min_periods=1).mean()
df['CO_MA168'] = df['CO(GT)'].rolling(window=168, min_periods=1).mean()
df['CO_std24'] = df['CO(GT)'].rolling(window=24, min_periods=1).std()
df['T_MA24'] = df['T'].rolling(window=24, min_periods=1).mean()
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

---

## Exercício 6: Lags e Diferenças

Crie features temporais usando lags e diferenças.

**Crie as seguintes colunas no DataFrame `df`:**
1. `CO_lag1`: CO de 1 hora atrás
2. `CO_lag24`: CO de 24 horas atrás
3. `CO_diff1`: diferença de CO em relação a 1 hora atrás
4. `CO_diff24`: diferença de CO em relação a 24 horas atrás

**Dica:** Use `.shift()` para lags e `.diff()` para diferenças

```python
df['CO_lag1'] = df['CO(GT)'].shift(1)
df['CO_lag24'] = df['CO(GT)'].shift(24)
df['CO_diff1'] = df['CO(GT)'].diff(1)
df['CO_diff24'] = df['CO(GT)'].diff(24)
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

---

## Exercício 7: Decomposição de Séries Temporais

Decomponha a série temporal de CO em seus componentes.

**Passos:**
1. Crie `co_daily_clean`: resample de CO para diário (média) e remova NaN com `.dropna()`
2. Importe `seasonal_decompose` de `statsmodels.tsa.seasonal`
3. Faça a decomposição com:
   - `model='additive'`
   - `period=7` (sazonalidade semanal)
   - `extrapolate_trend='freq'`
4. Extraia os componentes: `trend`, `seasonal`, `residual`

```python
from statsmodels.tsa.seasonal import seasonal_decompose

co_daily_clean = df['CO(GT)'].resample('D').mean().dropna()
decomp = seasonal_decompose(co_daily_clean, model='additive', period=7, extrapolate_trend='freq')
trend = decomp.trend
seasonal = decomp.seasonal
residual = decomp.resid
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

---

## Exercício 8: Análise de Componentes

Analise os componentes da decomposição.

**Calcule:**
1. `trend_change_pct`: variação percentual da tendência do primeiro ao último valor (não-NaN)
   - Fórmula: `((último - primeiro) / primeiro) * 100`
2. `seasonal_amplitude`: amplitude sazonal (máximo - mínimo)
3. `residual_std`: desvio padrão do resíduo

**Dica:** Use `.dropna()` antes de pegar primeiro/último valor com `.iloc[0]` e `.iloc[-1]`

```python
trend_clean = trend.dropna()
trend_change_pct = ((trend_clean.iloc[-1] - trend_clean.iloc[0]) / trend_clean.iloc[0]) * 100
seasonal_amplitude = seasonal.max() - seasonal.min()
residual_std = residual.std()
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

---

## Exercício 9: Detecção de Anomalias

Detecte anomalias na série de CO usando z-score móvel.

**Passos:**
1. Calcule `rolling_mean`: média móvel de 24 horas de CO
2. Calcule `rolling_std`: desvio padrão móvel de 24 horas de CO
3. Calcule `zscore`: z-score móvel = `(CO - rolling_mean) / rolling_std`
4. Identifique anomalias: `is_anomaly` = 1 quando `|zscore| > 3`, 0 caso contrário
5. Conte quantas anomalias foram detectadas em `num_anomalies`

**Dica:** Use `min_periods=1` no rolling para evitar NaN no início

```python
rolling_mean = df['CO(GT)'].rolling(window=24, min_periods=1).mean()
rolling_std = df['CO(GT)'].rolling(window=24, min_periods=1).std()
zscore = (df['CO(GT)'] - rolling_mean) / rolling_std
is_anomaly = (zscore.abs() > 3).astype(int)
num_anomalies = int(is_anomaly.sum())
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

---

## Exercício 10: Insights Finais

Responda às seguintes perguntas com base nas análises anteriores.

**Calcule:**
1. `worst_hour`: hora do dia com maior média de CO (0-23)
2. `worst_day`: dia da semana com maior média de CO (0-6, onde 0=Segunda)
3. `weekend_vs_weekday_ratio`: razão entre média de CO fim de semana / dias úteis
4. `trend_direction`: string 'up' se tendência crescente (trend_change_pct > 5), 'down' se decrescente (< -5), 'stable' caso contrário

**Dica:** Use os valores calculados anteriormente (co_mean_by_hour, co_mean_weekday, co_mean_weekend, trend_change_pct)

```python
worst_hour = int(co_mean_by_hour.idxmax())
worst_day = int(df.groupby('dayofweek')['CO(GT)'].mean().idxmax())
weekend_vs_weekday_ratio = co_mean_weekend / co_mean_weekday

if trend_change_pct > 5:
    trend_direction = 'up'
elif trend_change_pct < -5:
    trend_direction = 'down'
else:
    trend_direction = 'stable'
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

---

## Resumo:
- DatetimeIndex e manipulação de índices temporais
- Extração de componentes temporais
- Agregações por período
- Resample para diferentes frequências
- Rolling windows (médias móveis)
- Lags e diferenças
- Decomposição de séries temporais
- Análise de componentes
- Detecção de anomalias
- Insights de séries temporais



Tests for q1 (points: 8)
```
>>> assert isinstance(df.index, pd.DatetimeIndex), 'O índice deve ser DatetimeIndex'
>>> assert df.index.is_monotonic_increasing, 'O índice deve estar ordenado'
>>> assert not df.index.has_duplicates, 'O índice não deve ter duplicatas'
>>> assert len(df) > 9000, 'DataFrame deve ter mais de 9000 registros'
```


Tests for q2 (points: 6)
```
>>> assert 'hour' in df.columns, 'Coluna hour deve existir'
>>> assert 'dayofweek' in df.columns, 'Coluna dayofweek deve existir'
>>> assert 'month' in df.columns, 'Coluna month deve existir'
>>> assert 'is_weekend' in df.columns, 'Coluna is_weekend deve existir'
>>> assert 'is_rush_hour' in df.columns, 'Coluna is_rush_hour deve existir'
>>> assert df['hour'].min() >= 0 and df['hour'].max() <= 23, 'hour deve estar entre 0 e 23'
>>> assert df['dayofweek'].min() >= 0 and df['dayofweek'].max() <= 6, 'dayofweek deve estar entre 0 e 6'
>>> assert set(df['is_weekend'].unique()).issubset({0, 1}), 'is_weekend deve ser 0 ou 1'
```


Tests for q3 (points: 8)
```
>>> assert isinstance(co_mean_weekday, (int, float, np.number)), 'co_mean_weekday deve ser número'
>>> assert isinstance(co_mean_weekend, (int, float, np.number)), 'co_mean_weekend deve ser número'
>>> assert isinstance(co_mean_rushhour, (int, float, np.number)), 'co_mean_rushhour deve ser número'
>>> assert isinstance(co_mean_by_hour, pd.Series), 'co_mean_by_hour deve ser Series'
>>> assert len(co_mean_by_hour) == 24, 'co_mean_by_hour deve ter 24 valores'
>>> assert co_mean_weekday > 0, 'co_mean_weekday deve ser positivo'
```


Tests for q4 (points: 8)
```
>>> assert isinstance(co_daily, pd.Series), 'co_daily deve ser Series'
>>> assert isinstance(co_weekly, pd.Series), 'co_weekly deve ser Series'
>>> assert isinstance(co_daily_max, pd.Series), 'co_daily_max deve ser Series'
>>> assert isinstance(co_6h, pd.Series), 'co_6h deve ser Series'
>>> assert len(co_daily) > 300, 'co_daily deve ter mais de 300 registros'
>>> assert len(co_weekly) < len(co_daily), 'co_weekly deve ter menos registros que co_daily'
>>> assert len(co_6h) > len(co_daily), 'co_6h deve ter mais registros que co_daily'
```


Tests for q5 (points: 8)
```
>>> assert 'CO_MA24' in df.columns, 'Coluna CO_MA24 deve existir'
>>> assert 'CO_MA168' in df.columns, 'Coluna CO_MA168 deve existir'
>>> assert 'CO_std24' in df.columns, 'Coluna CO_std24 deve existir'
>>> assert 'T_MA24' in df.columns, 'Coluna T_MA24 deve existir'
>>> assert df['CO_MA24'].notna().sum() > 0, 'CO_MA24 deve ter valores não-NaN'
>>> assert df['CO_std24'].notna().sum() > 0, 'CO_std24 deve ter valores não-NaN'
```


Tests for q6 (points: 8)
```
>>> assert 'CO_lag1' in df.columns, 'Coluna CO_lag1 deve existir'
>>> assert 'CO_lag24' in df.columns, 'Coluna CO_lag24 deve existir'
>>> assert 'CO_diff1' in df.columns, 'Coluna CO_diff1 deve existir'
>>> assert 'CO_diff24' in df.columns, 'Coluna CO_diff24 deve existir'
>>> assert pd.isna(df['CO_lag1'].iloc[0]), 'Primeiro valor de CO_lag1 deve ser NaN'
>>> assert df['CO_diff1'].notna().sum() > 0, 'CO_diff1 deve ter valores não-NaN'
```


Tests for q7 (points: 8)
```
>>> assert isinstance(co_daily_clean, pd.Series), 'co_daily_clean deve ser Series'
>>> assert len(co_daily_clean) > 300, 'co_daily_clean deve ter mais de 300 dias'
>>> assert co_daily_clean.isna().sum() == 0, 'co_daily_clean não deve ter NaN'
>>> assert isinstance(trend, pd.Series), 'trend deve ser Series'
>>> assert isinstance(seasonal, pd.Series), 'seasonal deve ser Series'
>>> assert isinstance(residual, pd.Series), 'residual deve ser Series'
>>> assert len(trend) == len(co_daily_clean), 'trend deve ter mesmo tamanho que co_daily_clean'
```


Tests for q8 (points: 6)
```
>>> assert isinstance(trend_change_pct, (int, float, np.number)), 'trend_change_pct deve ser número'
>>> assert isinstance(seasonal_amplitude, (int, float, np.number)), 'seasonal_amplitude deve ser número'
>>> assert isinstance(residual_std, (int, float, np.number)), 'residual_std deve ser número'
>>> assert -100 < trend_change_pct < 100, 'trend_change_pct deve estar entre -100 e 100'
>>> assert seasonal_amplitude > 0, 'seasonal_amplitude deve ser positivo'
>>> assert residual_std > 0, 'residual_std deve ser positivo'
```


Tests for q9 (points: 8)
```
>>> assert isinstance(rolling_mean, pd.Series), 'rolling_mean deve ser Series'
>>> assert isinstance(rolling_std, pd.Series), 'rolling_std deve ser Series'
>>> assert isinstance(zscore, pd.Series), 'zscore deve ser Series'
>>> assert isinstance(is_anomaly, pd.Series), 'is_anomaly deve ser Series'
>>> assert isinstance(num_anomalies, (int, np.integer)), 'num_anomalies deve ser inteiro'
>>> assert set(is_anomaly.dropna().unique()).issubset({0, 1}), 'is_anomaly deve ser 0 ou 1'
>>> assert num_anomalies >= 0, 'num_anomalies deve ser não-negativo'
```


Tests for q10 (points: 8)
```
>>> assert isinstance(worst_hour, (int, np.integer)), 'worst_hour deve ser inteiro'
>>> assert isinstance(worst_day, (int, np.integer)), 'worst_day deve ser inteiro'
>>> assert isinstance(weekend_vs_weekday_ratio, (int, float, np.number)), 'weekend_vs_weekday_ratio deve ser número'
>>> assert isinstance(trend_direction, str), 'trend_direction deve ser string'
>>> assert 0 <= worst_hour <= 23, 'worst_hour deve estar entre 0 e 23'
>>> assert 0 <= worst_day <= 6, 'worst_day deve estar entre 0 e 6'
>>> assert trend_direction in ['up', 'down', 'stable'], 'trend_direction deve ser up, down ou stable'
```
